Effects of blood meal source and seasonality on reproductive traits of Culex quinquefasciatus (Diptera: Culicidae)

Author response:

The following is the authors’ response to the previous reviews.

We have carefully addressed all the reviewers' suggestions, and detailed responses are provided at the end of this letter. In summary:

• We conducted two additional replicates of the study to obtain more robust and reliable data.

• The Introduction has been revised for greater clarity and conciseness.

• The Results section was shortened and reorganized to highlight the key findings more effectively.

• The Discussion was modified according to the reviewers' suggestions, with a focus on reorganization and conciseness.

We hope you find this revised version of the manuscript satisfactory.

Reviewer #1 (Public Review):

Summary:

This study examines the role of host blood meal source, temperature, and photoperiod on the reproductive traits of Cx. quinquefasciatus, an important vector of numerous pathogens of medical importance. The host use pattern of Cx. quinquefasciatus is interesting in that it feeds on birds during spring and shifts to feeding on mammals towards fall. Various hypotheses have been proposed to explain the seasonal shift in host use in this species but have provided limited evidence. This study examines whether the shifting of host classes from birds to mammals towards autumn offers any reproductive advantages to Cx. quinquefasciatus in terms of enhanced fecundity, fertility, and hatchability of the offspring. The authors found no evidence of this, suggesting that alternate mechanisms may drive the seasonal shift in host use in Cx. quinquefasciatus.

Strengths:

Host blood meal source, temperature, and photoperiod were all examined together.

Weaknesses:

The study was conducted in laboratory conditions with a local population of Cx. quinquefasciatus from Argentina. I'm not sure if there is any evidence for a seasonal shift in the host use pattern in Cx. quinquefasciatus populations from the southern latitudes.

Comments on the revision:

Overall, I am not quite convinced about the possible shift in host use in the Argentinian populations of Cx. quinquefasciatus. The evidence from the papers that the authors cite is not strong enough to derive this conclusion. Therefore, I think that the introduction and discussion parts where they talk about host shift in Cx. quinquefasciatus should be removed completely as it misleads the readers. I suggest limiting the manuscript to talking only about the effects of blood meal source and seasonality on the reproductive outcomes of Cx. quinquefasciatus.

As mentioned in the previous revision, we agree on the reviewer observation about the lack of evidence on seasonal shift in the host use pattern in Cx. quinquefasciatus populations from Argentina. We include this topic in the discussion.

Additionally, we also added a paragraph in the discussion section to include the limitations of our study and conclusions. One of them is the fact that our results are based on controlled conditions experiments. Future studies are needed to elucidate if the same trend is found in the field.

Reviewer #1 (Recommendations for the authors):

Abstract

Line 73: shift in feeding behavior

Accepted as suggested.

Discussion

Line 258: addressed that Accepted as suggested.

Line 263: blood is nutritionally richer

Accepted as suggested.

Reviewer #2 (Public Review):

Summary:

Conceptually, this study is interesting and is the first attempt to account for the potentially interactive effects of seasonality and blood source on mosquito fitness, which the authors frame as a possible explanation for previously observed host-switching of Culex quinquefasciatus from birds to mammals in the fall. The authors hypothesize that if changes in fitness by blood source change between seasons, higher fitness on birds in the summer and on mammals in the autumn could drive observed host switching. To test this, the authors fed individuals from a colony of Cx. quinquefasciatus on chickens (bird model) and mice (mammal model) and subjected each of these two groups to two different environmental conditions reflecting the high and low temperatures and photoperiod experienced in summer and autumn in Córdoba, Argentina (aka seasonality). They measured fecundity, fertility, and hatchability over two gonotrophic cycles. The authors then used a generalized linear model to evaluate the impact of host species, seasonality, and gonotrophic cycle on fecundity, fertility, and hatchability. The authors were trying to test their hypothesis by determining whether there was an interactive effect of season and host species on mosquito fitness. This is an interesting hypothesis; if it had been supported, it would provide support for a new mechanism driving host switching. While the authors did report an interactive impact of seasonality and host species, the directionality of the effect was the opposite from that hypothesized. The authors have done a very good job of addressing many of the reviewer concerns, with several exception that continue to cause concern about the conclusions of the study.

Strengths:

(1) Using a combination of laboratory feedings and incubators to simulate seasonal environmental conditions is a good, controlled way to assess the potentially interactive impact of host species and seasonality on the fitness of Culex quinquefasciatus in the lab.

(2) The driving hypothesis is an interesting and creative way to think about a potential driver of host switching observed in the field.

(3) The manuscript has become a lot clearer and easier to read with the revisions - thank you to the authors for working hard to make many of the suggested changes.

Weaknesses:

(1) The authors have decided not to follow the suggestion of conducting experimental replicates of the study. This is understandable given the significant investment of resources and time necessary, however, it leaves the study lacking support. Experimental replication is an important feature of a strong study and helps to provide confidence that the observed patterns are real and replicable. Without replication, I continue to lack confidence in the conclusions of the study.

We included replicates as suggested.

(2) The authors have included some additional discussion about the counterintuitive nature of their results, but the paragraph discussing this in the discussion was confusing. I believe that this should be revised. This is a key point of the paper and needs to be clear to the reader.

Revised as suggested.

(3) There should be more discussion of the host switching observed in the two studies conducted in Argentina referenced by the authors. Since host switching is the foundation for the hypothesis tested in this paper, it is important to fully explain what is currently known in Argentina.

Accepted as suggested.

(4) In some cases, the explanations of referenced papers are not entirely accurate. For example, when referencing Erram et al 2022, I think the authors misrepresented the paper's discussion regarding pre-diuresis- Erram et al. are suggesting that pre-diuresis might be the mechanism by which C. furens compensates for the lower nutritional value of avian blood, leading to no significant difference between avian/mammal blood on fecundity/fertility (rather than leading to higher fecundity on birds, as stated in this manuscript). The study performed by Erram et al. also didn't prove this phenomenon, they just suggest it as a possible mechanism to explain their results, so that should be made clear when referencing the paper.

Changed as suggested.

(5) In some cases, the conclusions continue to be too strongly worded for the evidence available. For example, lines 322-324: I don't think the data is sufficient to conclude that a different physiological state is induced, nor that they are required to feed on a blood source that results in higher fitness.

Redaction was modified as suggested to tight our discussion with results.

(6) There is limited mention of the caveat that this experiment performed with simulated seasonality that does not perfectly replicate seasonality in the field. I think this caveat should be discussed in the discussion (e.g. that humidity is held constant).

This topic is now included in the discussion as suggested.

Reviewer #2 (Recommendations for the authors):

59-60: These terms should end with -phagic instead of -philic. These papers study blood feeding patterns, not preference. I understand that the Janssen papers calls it "mammalophilic" in their title, but this was an incorrect use of the term in their paper. There are some review papers that explain the difference in this terminology if it's helpful.

Accepted as suggested.

73: edit to "in" feeding behavior

Accepted as suggested.

77-78: Given that the premise of your study is based on the phenomenon of host switching, I suggest that you expand your discussion of these two papers. What did they observe? Which hosts did they switch from / to and how dramatic was the shift?

Accepted as suggested.

79: replace acknowledged with experienced

Accepted as suggested.

79-80: the way that this is written is misleading. It suggests that Spinsanti showed that seasonal variation in SLEV could be attributed to a host shift, which isn't true. This citation should come before the comma and then you should use more cautious language in the second half. E.g which MIGHT be possible to attribute to ....

Accepted as suggested.

80-82: this is not convincing. Even if the Robin isn't in Argentina, Argentina does have migrating birds, so couldn't this be the case for other species of birds? Do any of the birds observed in previous blood meal analyses in Argentina migrate? If so, couldn't this hypothesis indeed play a role?

A paragraph about this topic was added to the discussion as suggested.

90: hypotheses for what? The fall peak in cases? Or host switching?

Changed to be clearer.

98: where was this mentioned before? I think "as mentioned before" can be removed.

Accepted as suggested.

101: edit to "whether an interaction effect exists"

Accepted as suggested.

104: edit to "We hypothesize that..."

Accepted as suggested.

106: reported host USE changes, not host PREFERENCE changes, right?

All the terminology was change to host pattern and not preference to avoid confusion.

200: Briefly reading Carsey and Harden, it looks like the methodology was developed for social science. Is there anything you can cite to show this applied to other types of data? If not, I think this requires more explanation in your MS.

This was removed as replicates were included.

237-239: I think it is best not to make a definitive statement about greater/higher if it isn't statistically significant; I suggest modifying the sentences to state that the differences you are listing were not significantly different up front rather than at the end, otherwise if people aren't reading carefully, they may get the wrong impression.

Accepted as suggested.

245: you only use the term MS-I once before and I forgot what it meant since it wasn't repeated, so I had to search back through with command-F. I suggest writing this out rather than using the acronym.

Accepted as suggested.

249: edit to: "an interaction exists between the effect of..."

Accepted as suggested.

253-254: greater compared to what?

Change for clearness. 258-260: edit for grammar

Accepted as suggested.

260-262: edit for grammar; e.g. "However, this assumption lacks solid evidence; there is a scarcity of studies regarding nutritional quality of avian blood and its impact on mosquito fitness."

Accepted as suggested.

263: edit: blood is nutritionally...

Accepted as suggested.

264-267: This doesn't sound like an accurate interpretation of what the paper suggests regarding pre-diuresis in their discussion - they are suggesting that pre-diuresis might be the mechanism by which C. furens compensates for the lower nutritional value of avian blood, leading to no significant difference between avian/mammal blood on fecundity/fertility. They also don't show this, they just suggest it as a possible mechanism to explain their results.

This topic was removed given the restructuring of discussion.

253-269: You should tie this paragraph back to your results to explicitly compare/contrast your findings with the previous literature.

Accepted as suggested.

270-282: This paragraph would be a good place to explain the caveat of working in the laboratory - for example, humidity was the same across the two seasons which I'm guessing isn't the case in the field in Argentina. You can discuss what aspects of laboratory season simulation do not accurately replicate field conditions and how this can impact your findings. You said in your response to the reviewers that you weren't interested in measuring other variables (which is fair, and not expected!), but the beauty of the discussion section is to be able to think about how your experimental design might impact your results - one possibility is that your season simulation may not have produced the results produced by true seasonal shifts.

Accepted as suggested.

279-281: You say your experiment was conducted within the optimal range, which would suggest that both summer and autumn were within that range, but then you only talk about summer as optimal in the following sentence.

Changed for clearness.

281-282: You should clarify this sentence - state what the interaction has an effect on.

Accepted as suggested.

283-291: I appreciate that your discussion now acknowledges the small sample size and the questions that remain unanswered due to the results being opposite to that of the hypothesis, but this paragraph lacks some details and in places doesn't make sense.

I think you need to emphasize which groups had small sample size and which conclusions that might impact. I also think you need to explain why the sample size was substantially smaller for some groups (e.g. did they refuse to feed on the mouse in the autumn?). I appreciate that sample sizes are hard to keep high across many groups and two gonotrophic periods, but unfortunately, that is why fitness experiments are so hard to do and by their nature, take a long time. I understand that other papers have even lower sample size, but I was not asked to review those papers and would have had the same critique of them. I don't believe that creating simulated data via a Monte Carlo approach can make up for generating real data. As I understand it from your explanation, you are parametrizing the Monte Carlo simulations with your original data, which was small to begin with for autumn mouse. Using this simulation doesn't seem like a satisfactory replacement for an experimental replicate in my opinion. I maintain that at least a second replicate is necessary to see whether the patterns that you have observed hold.

The performing of a power analysis and addition of more replicates tried to solve the issue of sample size. More about this critic is added in the discussion. The simulation approach was totally removed.

Regarding the directionality of the interaction effect, I think this warrants more discussion. Lines 287-291 don't make sense to me. You suggest that feeding on birds in the autumn may confer a reproductive advantage when conditions are more challenging. But then why wouldn't they preferentially feed on birds in the autumn, rather than mammals? I suggest rewriting this paragraph to make it clearer.

Accepted as suggested.

297: earlier mentioned treatments? Do you mean compared to the first gonotrophic cycle? This isn't clear.

Changed for clearness.

302-303: Did you clarify whether you are allowed to reference unpublished data in eLife?

This was removed to follow the guidelines of eLife.

316-317: "it becomes apparent" sounds awkward, I suggest rewording and also explaining how this conclusion was made.

Accepted as suggested.

322-324: I think that this statement is too strongly worded. I don't think your data is sufficient to conclude that a different physiological state is induced, nor that they are required to feed on a blood source that results in higher fitness. Please modify this and make your conclusions more cautious and closely linked to what you actually demonstrated.

Accepted as suggested.

325: change will perform to would have

Accepted as suggested.

326: add to the sentence: "and vice versa in the summer"

Accepted as suggested.

330: possible explanations, not explaining scenarios.

Accepted as suggested.

517: I think you should repeat the abbreviation definitions in the caption to make it easier for readers, otherwise they have to flip back and forth which can be difficult depending on formatting.

Accepted as suggested.

In general, I think that your captions need more information. I think the best captions explain the figure relatively thoroughly such that the reader can look at the figure and caption and understand without reading the paper in depth. (e.g. the statistical test used).

Data availability: The eLife author instructions do say that data must be made available, so there should be a statement on data availability in your MS. I also suggest you make the code available.

Accepted as suggested.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Host selection by mosquitoes is a keystone to understanding viral circulation and predicting future infection outbreaks. Culex mosquitoes frequently feed on birds during spring and early summer, shifting into mammals towards late summer and autumn. This host switch may be due to changes in mosquito fitness. The aim of this study was to assess if the interaction effect of blood meal source and seasonality may influence reproductive traits of Culex quinquefasciatus mosquitoes. For this purpose, Cx. quinquefasciatus mosquitoes were reared in simulated summer and autumn conditions and fed on two different hosts, chickens and mice, in a factorial design. Fecundity, fertility, and hatchability during two consecutive gonotrophic cycles were estimated. We found greater fecundity and fertility for mosquitoes fed upon birds than mammals. Fecundity and fertility increased in autumn for chicken-fed mosquitoes, whereas they decreased for mouse-fed mosquitoes. These traits decreased in the second gonotrophic cycle for mouse-fed mosquitoes, whereas they did not vary between cycles for chicken-fed mosquitoes. The effect of hatchability on fertility was rather limited. These results indicate a statistically significant interaction effect of blood meal source and seasonality on fecundity and fertility. However, the pattern was opposite in relation to our hypothesis, suggesting that further studies are needed to confirm and expand our knowledge about mosquito biology and its relationship with seasonal host use shifting.

Revised version: Host selection by mosquitoes is a keystone to understanding viral circulation and predicting future infection outbreaks. Culex mosquitoes frequently feed on birds during spring and early summer, shifting into mammals towards late summer and autumn. This host switch may be due to changes in mosquito fitness. The aim of this study was to assess if the interaction effect of blood meal source and seasonality may influence reproductive traits of Culex quinquefasciatus mosquitoes. For this purpose, Cx. quinquefasciatus mosquitoes were reared in simulated summer and autumn conditions and fed on two different hosts, chickens and mice, in a factorial design. Fecundity, fertility, and hatchability during two consecutive gonotrophic cycles were estimated. We found greater fecundity and fertility for mosquitoes fed upon birds than mammals. Fecundity and fertility did not vary between seasons for chicken-fed mosquitoes, whereas in autumn they decreased for mouse-fed mosquitoes. These traits decreased in the second gonotrophic cycle for mouse-fed mosquitoes, whereas they did not vary between cycles for chicken-fed mosquitoes. There was no statistically significant variation of hatchability among treatments. These results indicate a statistically significant interaction effect of blood meal source and seasonality on fecundity and fertility. However, the pattern was opposite in relation to our hypothesis, suggesting that further studies are needed to confirm and expand our knowledge about mosquito biology and its relationship with seasonal host use shifting.

Introduction

Earlier version: The southern house mosquito, Culex quinquefasciatus Say (Diptera: Culicidae), is a worldwide distributed vector of several pathogens, such as Wuchereria bancrofti Cobbold (Kasili et al., 2009), Dirofilaria immitis Leidy (Labarthe et al., 1998) and avian malaria (Farajollahi et al., 2011). Besides, this mosquito is also one of the main species responsible for the transmission of the arboviruses West Nile virus (WNV) and Saint Louis encephalitis virus (SLEV), among others (Farajollahi et al., 2011). Because of this, this species is of critical concern due to its impact on both public and veterinary health.

Revised version: The southern house mosquito, Culex quinquefasciatus Say (Diptera: Culicidae), is a worldwide distributed vector of numerous pathogens, including Wuchereria bancrofti Cobbold (Kasili et al., 2009), Dirofilaria immitis Leidy (Labarthe et al., 1998) and avian malaria (Farajollahi et al., 2011). Besides, this mosquito is also one of the main species responsible for the transmission of the arboviruses West Nile (WNV) and St. Louis encephalitis (SLEV) viruses (Farajollahi et al., 2011). Because this, this species is of critical concern due to its impact on both public and veterinary health.

Earlier version: Since arboviruses are transmitted by the bite of infected mosquitoes between vertebrate hosts, activity of these pathogens is determined, at least in part, by the blood feeding habits of these vectors. Host preference of Culex mosquitoes is variable and it can be classified in four categories: (1) mammophilic (Muturi et al., 2008), (2) ornithophilic (Jansen et al., 2009), (3) herpetophilic (Janssen et al., 2015) and (4) generalist (Molaei et al., 2007). Additionally, some species may experience a host switch in its feeding habits, from avian to mammal hosts (Hancock & Camp, 2022). These species are of greater concern from an epidemiological standpoint because they can act as bridge vectors in the transmission of some arboviruses, such as SLEV and WNV (Kilpatrick et al., 2005).

Revised version: Since arboviruses are transmitted by the bite of infected mosquitoes between vertebrate hosts, activity of these pathogens is determined, at least in part, by the blood feeding habits of these vectors. According to its host feeding pattern range, Culex mosquitoes can be regarded as specialists or generalists. Specialist species include mosquitoes feeding primarily on mammals (Muturi et al., 2008), birds (Jansen et al., 2009), reptiles and amphibians (Janssen et al., 2015). Generalists are mosquitoes fed on the most abundant available hosts (Molaei et al., 2007). Additionally, some species could experience a seasonal shift in its feeding habits, from avian to mammal hosts (Edman & Taylor, 1968; Hancock & Camp, 2022), being of greater concern from an epidemiological perspective given that they can act as bridge vectors in the transmission of some arboviruses, such as SLEV and WNV (Kilpatrick et al., 2006).

Added in the revised version: West Nile virus (WNV) and St. Louis encephalitis virus (SLEV) are zoonotic pathogens maintained in nature through an enzootic network involving transmission among viremic birds and ornithophagic mosquitoes (Kilpatrick et al., 2006). However, during late summer and early autumn, these viruses increase their activity in humans and other mammals, which has led to repeated epidemics in the US (Kilpatrick et al., 2006) and South America (Spinsanti et al., 2008). Likewise, several Culex species, including Cx. quinquefasciatus, have exhibited a seasonal shift in their feeding behaviour, feeding primarily on birds in spring and early summer and increasing the proportion of mammal blood meals in late summer and autumn (Thiemann et al., 2011). This seasonal shift suggests that viral activity might spill over from birds to mammals as a result of the mosquito host switch (Edman & Taylor, 1968; Kilpatrick et al., 2006).

Added in the revised version: Two main contrasting hypotheses have been proposed to explain why primarily ornithophagic mosquitoes increase their feeding on mammals later in the season. The migration hypothesis, proposed by Kilpatrick et al. (2006) for Cx. pipiens, suggests that this shift in feeding pattern is related to the abundance of the preferred avian host, the American robin (Turdus migratorius L.). This hypothesis is based on their findings that as robin densities decrease, the proportion of mammalian blood meals in Cx. pipiens populations increase. However, this hypothesis has several limitations: (i) multiple studies have shown that American robins are not always the preferred host for all Cx. pipiens populations (Patrican et al., 2007; Apperson et al., 2002); (ii) in some areas (e.g., Tennessee) with documented seasonal host shifts, American robins are present year-round (Savage et al., 2007); and (iii) this explanation applies only to mosquitoes whose primary host is the American robin, a species found only in the Northern Hemisphere, or another migratory bird. Therefore, it is not applicable to southern regions, such as Argentina, where SLEV seasonality is documented (Spinsanti et al., 2008) but American robins are absent. Because this, it seems that the hypothesis proposed by Kilpatrick et al. (2006) might be habitat dependent (Caranci, 2010). On the other hand, the defensive behaviour hypothesis proposed by Burkett-Cadena et al. (2011), suggests that host breeding cycles drive the host shift of mosquitoes. According to this hypothesis, in hosts with parental care, during reproductive seasons there is a greater investment of energy in assuring the offspring survivorship, consequently producing an increase in susceptibility of being bitten by mosquitoes as a result of a decrease in defensive behaviours. This event leads to the detection of peaks of host use during periods of reproductive investment, in summer for birds and autumn for mammals (Burkett-Cadena et al., 2011). While this hypothesis was proposed as an alternative to address the limitations of the migration hypothesis, it also has drawbacks. It is based on a limited range of host species and assumes a distinct seasonal phenology, which is not present for all mosquito hosts. Because this, this hypothesis seems to be host dependent.

Removed in the revised version: Culex quinquefasciatus is generally considered an ornithophilic species in several areas around the globe (Takken & Verhulst, 2013). Despite its ornithophilic nature, this and other Culex species sometimes feed primarily on mammal hosts, including humans, in certain situations. Various studies have pointed out that in autumn there is an increased activity of SLEV and WNV in human populations (Kilpatrick et al., 2006; Spinsanti et al., 2008), suggesting that viral activity spills over to mammals from birds as a result of the mosquito host switch (Edman & Taylor, 1968; Kilpatrick et al., 2006). A few hypotheses have attempted to explain why primarily ornithophilic mosquitoes also feed upon mammals in some contexts. One of the most widely accepted hypothesis states that the shift on feeding behaviour of Culex mosquitoes is a consequence of the autumn migration of American robin (Turdus migratorius), the principal host in the northeastern US, generating an opportunistic change into mammal hosts (Kilpatrick et al., 2006). In South America in general and in Argentina in particular, studies regarding host switch are scarce. Nevertheless, a few limited studies have indicated a seasonal variation throughout the seasons (Stein et al., 2013; Beranek, 2018). Despite that, Argentina has acknowledged seasonal variation in the activity of SLEV in humans, which can be attributed to variations in mosquito host-feeding pattern (Spinsanti et al., 2008). In Southern latitudes, species such as T. migratorius are lacking, hence the “migration hypothesis” is very unlikely for many austral areas, being restricted only for some regions of the US.

Removed in the revised version: A second hypothesis proposed by Burkett-Cadena et al. (2011) suggests that host breeding cycles drive the host shift of mosquitoes. According to this hypothesis, in hosts with parental care, during reproductive seasons there is a greater investment of energy in assuring the offspring survivorship, consequently producing an increase in susceptibility of being bitten by mosquitoes as a result of a decrease in defensive behaviours. This event leads to the detection of peaks of host use during periods of reproductive investment, in summer for birds and autumn for mammals (Burkett-Cadena et al., 2011).

Earlier version: While some other minor hypotheses, unsupported by field data, have been proposed, none of them have been specifically focused on vector biology. Numerous biological factors, including stress, metabolic rate, and the source of a blood meal, as well as environmental variables such as temperature and photoperiod, exert an influence on mosquito physiology that might affect various reproductive traits such as fecundity, development rate or survivorship. This set of factors gives rise to new nutritional requirements that, ultimately, may lead to seasonal variation in host choice by mosquitoes (Ciota et al., 2014; Costanzo et al., 2015; Gervasi et al., 2016; Yan et al. 2017, 2018). In the literature, several studies have evaluated the impact of these biological and environmental variables on mosquito reproduction, as mentioned before. However, due to the complex nature of mosquito biology, there are likely multiple interactions among variables that give rise to new responses that may not be observed individually. Therefore, the aim of the present study was to assess whether exists an interaction effect between the source of the host blood meal and seasonality (in terms of temperature and photoperiod) on three reproductive traits of Culex quinquefasciatus mosquitoes, fecundity, fertility, and hatchability. Our hypothesis states that the interaction between these two variables influence the reproductive outcomes, potentially leading to a seasonal shift in host selection, driven by a reproductive advantage. Considering the reported host preference changes in Cx. quinquefasciatus, in autumn we expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes after feeding on a mammal host compared to an avian host, and the opposite relationship in summer.

Revised version: Explanations for the seasonal shift in mosquito host feeding focused on vector biology are largely lacking in the literature. Numerous biological factors, such as stress, metabolic rate, and blood meal source, along with environmental variables like temperature and photoperiod, influence mosquito physiology and can affect reproductive traits such as fecundity, development rate, and survivorship. These factors can give rise to new nutritional requirements that may ultimately lead to seasonal variations in host selection by mosquitoes (Ciota et al., 2014; Costanzo et al., 2015; Gervasi et al., 2016; Yan et al., 2017, 2018). While several studies have examined the effects of these biological and environmental variables on mosquito reproduction, the complex nature of mosquito biology suggests that multiple interactions among these variables may produce novel responses that are not apparent when considered individually. The aim of this study was to assess whether there is an interaction between the source of the host blood meal and seasonality (defined by temperature and photoperiod) on three reproductive traits of Culex quinquefasciatus mosquitoes: fecundity, fertility, and hatchability. We hypothesize that the interaction between these two variables influence the reproductive outcomes, potentially leading to a seasonal shift in host selection, driven by a reproductive advantage. Given the reported seasonal changes in host use by Cx. quinquefasciatus, in autumn we expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes after feeding on a mammal host compared to an avian host, and the opposite trend in summer.

Materials and methods

Establishment and maintenance of mosquitoes

Earlier version: Egg rafts of Culex quinquefasciatus were collected from a drainage ditch at Universidad Nacional de Córdoba Campus, Córdoba city, in February 2021. Each raft was individually maintained in plastic containers with one liter of distilled water. The hatched larvae were fed with 100 mg of liver powder three times per week until pupation. Pupae were then transferred to plastic emerging cages (21 cm x 12 cm) covered with a tulle-like fabric, containing distilled water but no food. Adults emerging from each raft were identified morphologically (Darsie, 1985) and molecularly (Smith & Fonseca, 2004) to ensure they corresponded to Cx. quinquefasciatus, since Cx. pipiens and its hybrids are also present in Córdoba (Branda et al., 2021). All adults were reared in a cardboard cage of 22.5 liters (28 cm x 36.5 cm) and provided ad libitum with a 10% sugar solution soaked in cotton pads placed on plastic cups. For long-term maintenance of the colony, 24-hour-starved mosquitoes were offered a blood meal from a restrained chicken twice a month. Four days after feeding, a plastic container with distilled water was placed inside the cage to allow engorged females to lay egg rafts. Batches of egg were collected and transferred to plastic containers (30 cm x 25 cm x 7 cm) in a proportion of 3 rafts per container, filled with three liters of distilled water. The hatched larvae were also fed with liver powder at the same proportion described above. Pupae were transferred to emerging cages, and adult mosquitoes were placed in the final cardboard cage. The colony has been maintained for over twenty generations in the Insectary of the Instituto de Virología “Dr. J.M. Vanella” (InViV). Room conditions were controlled at 28°C, with a 12L:12D photoperiod and 70% relative humidity.

Revised version: Egg rafts of Culex quinquefasciatus were collected from a drainage ditch at Universidad Nacional de Córdoba Campus, Córdoba city, in February 2021. Each raft was individually maintained in plastic containers with one liter of distilled water. The hatched larvae were fed with 100 mg of liver powder three times per week until pupation. Pupae were then transferred to plastic emerging cages (21 cm x 12 cm) covered with a tulle-like fabric, containing distilled water but no food. Adults emerging from each raft were identified morphologically (Darsie, 1985) and molecularly (Smith & Fonseca, 2004) to ensure they corresponded to Cx. quinquefasciatus, since Cx. pipiens and its hybrids coexist sympatrically in Córdoba city (Branda et al., 2021). All adults were reared in a cardboard cage of 22.5 liters (28 cm x 36.5 cm) and provided ad libitum with a 10% sugar solution soaked in cotton pads placed on plastic cups. For long-term maintenance of the colony, 24-hour-starved mosquitoes were offered a blood meal from a restrained chicken twice a month. Four days after feeding, a plastic container with distilled water was placed inside the cage to allow engorged females to lay egg rafts. Batches of egg were collected and transferred to plastic containers (30 cm x 25 cm x 7 cm) in a proportion of three rafts per container, filled with three liters of distilled water. The hatched larvae were also fed with liver powder at the same proportion described above. Pupae were transferred to emerging cages, and adult mosquitoes were placed in the final cardboard cage. The colony has been maintained for several generations in the Insectary of the Instituto de Virología “Dr. J.M. Vanella” (InViV). Room conditions were controlled at 28 ± 1 °C, with a 12L:12D photoperiod and 70% relative humidity.

Experimental design

Blood source

For experimental trials, avian hosts (live chicks of the species Gallus gallus) and mammalian hosts (live mice of the species Mus musculus, strain C57BL/6) were used to evaluate the effect of blood meal source on fecundity, fertility, and hatchability. Chicks were generously donated by the Bartolucci poultry farm (Córdoba, Argentina), while mice were commercially obtained from the Instituto de Investigación Médica Mercedes y Martín Ferreyra (CONICET - Universidad Nacional de Córdoba). In each trial, 24-hour-starved adult female mosquitoes were provided with a blood meal from restrained chicks or mice. Vertebrate hosts were offered to mosquitoes one hour before the lights were turned off and were kept for 3 hours during 2 consecutive gonotrophic cycles.

Earlier version: The experimental use of animals (G. gallus and M. musculus) was approved by the ethical committee at Facultad de Ciencias Médicas, Universidad Nacional de Córdoba (FCM-UNC) in compliance with the legislation regarding the use of animals for experimental and other scientific purposes (accession code: CE-2022-00518476-UNC-SCT#FCM).

Revised version: The experimental use of animals (mosquitoes and vertebrates) was approved by the ethical committee at Facultad de Ciencias Médicas, Universidad Nacional de Córdoba (FCM-UNC) in compliance with the legislation regarding the use of animals for experimental and other scientific purposes (accession code: CE-2022-00518476-UNC-SCT#FCM).

Seasonality

To assess the effect of seasonality (photoperiod + temperature) on fecundity, fertility, and hatchability, a “typical” summer and autumn day from Córdoba city was simulated in an incubator where mosquitoes were housed throughout the entire experiment. Summer conditions were as follows: T°min = 22°C, T°max = 28°C, photoperiod = 14L:10D. Autumn conditions were characterized by: T°min = 16°C, T°max = 22°C, and photoperiod = 10L:14D. The humidity levels were maintained at 60-70% for both conditions. Data for simulated conditions were obtained from Climate Data website (https://es.climate-data.org/).

Feeding trial design

Earlier version: The interaction effect of blood source (bird or mammal) and seasonality (autumn and summer) was evaluated during two consecutive gonotrophic cycles (I and II). Eight experimental colonies were established using egg rafts collected from the maintenance colony: mammal-autumn-I (MA-I), mammal-summer-I (MS-I), bird-autumn-I (BA-I), bird-summer-I (BS-I), mammal-autumn-II (MA-II), mammal-summer-II (MS-II), bird-autumn-II (BA-II), and bird-summer-II (BS-II).

Revised version: The interaction effect of blood source (chicken and mouse) and seasonality (autumn and summer) was evaluated during two consecutive gonotrophic cycles (I and II). This combination produced eight colonies that were established using egg rafts collected from the maintenance colony: mouse:autumn-I, mouse:autumn-II, chicken:autumn-I, chicken:autumn-II, mouse:summer-I, mouse:summer-II, chicken:summer-I, chicken:summer-II. This set of 8 colonies was repeated in triplicate (1-3), generating a total of 24 treatments.

The experimental colonies were maintained under controlled conditions to ensure that all adults were of the same size (Kauffman et al., 2017). Adult mosquitoes were placed in cardboard cages of 8.3 liters (21 cm x 24 cm) and were provided with ad libitum access to a 10% sugar solution, which was soaked in cotton pads placed on plastic cups.

Earlier version: Feeding trials were conducted at two time points: 5 days post-emergence (first cycle) and 14 days post-emergence (second cycle). Following each blood meal, female mosquitoes were anesthetized using CO2 and classified as fully engorged (VI Sella’s stage), partially fed (I-V Sella’s stage), or unfed (I Sella’s stage), following the classification by Silva Santos et al. (2019). Fully engorged females were counted and separated into a separate cage to complete their gonotrophic cycle. Four days after feeding, a cup containing distilled water was placed inside the cage to facilitate oviposition. After each oviposition cycle, egg rafts were collected, counted, and transferred to a 12-well plate containing 3 mL of distilled water. They were then photographed to subsequently determine the number of eggs per raft. Rafts were maintained in each well until they hatched into L1 larvae, at which point they were preserved with 1 mL of 96% ethanol, and the number of L1 larvae per raft was counted.

Revised version: Feeding trials were conducted at two time points: 5 days post-emergence (first cycle) and 14 days post-emergence (second cycle). Following each blood meal, female mosquitoes were anesthetized using CO2 and classified as fully engorged (VI Sella’s stage), partially fed (I-V Sella’s stage), or unfed (I Sella’s stage), following the classification by Silva Santos et al. (2019). Fully engorged females were counted and separated into a separate cage to complete their gonotrophic cycle. Four days after feeding, a cup containing distilled water was placed inside the cage to facilitate oviposition. After each oviposition cycle, egg rafts were collected, counted, and transferred to a 12-well plate containing 4 mL of distilled water. They were then photographed to subsequently determine the number of eggs per raft. Rafts were maintained in each well until they hatched into L1 larvae, at which point they were preserved with 2 mL of 96% ethanol, and the number of L1 larvae per raft was counted.

Statistical analysis

The reproductive outputs measured for each experimental colony were fecundity, fertility, and hatchability. All analyses were performed using R Studio statistic software, v.4.2.1 (R Core Team, 2022).

Earlier version: Fecundity was defined as the number of eggs per raft and fertility as the number of L1 larvae hatched per raft. To evaluate the effect of blood meal source and seasonality on these two variables, a Generalized Linear Model (GLM; package MASS) with negative binomial error distribution and logarithmic link function was adjusted (Venables & Ripley, 2002). Fecundity and fertility served as the response variables, while the explanatory variables included blood source, seasonality, and gonotrophic cycle, all in a three-way interaction.

Revised version: Fecundity was defined as the number of eggs per raft and fertility as the number of L1 larvae hatched per raft. To evaluate the effect of blood meal source and seasonality on these two variables, a Generalized Linear Mixed Model (GLMM; lme4 package) with negative binomial error distribution and logarithmic link function was adjusted (Bates et al., 2015). Fecundity and fertility served as the response variables, while the fixed explanatory variables included blood source, seasonality, and gonotrophic cycle, all in a three-way interaction. Treatments (n = 24) were set as a random intercept.

Earlier version: Hatchability (= hatching rate) was defined as the ratio of the number of larvae to the number of eggs per raft. A GLM was applied using a quasipoisson error distribution and a logarithmic link function. The response variable was the number of larvae, while the explanatory variables included blood source, seasonality and gonotrophic cycle as interaction effects. Additionally, the logarithm of the number of eggs was incorporated as an offset in the model.

Revised version: Hatchability (= hatching rate) was defined as the ratio of the number of larvae to the number of eggs per raft. A Generalized Linear Model (GLM; MASS package) was applied using a quasipoisson error distribution and a logarithmic link function (Venables et al., 2002). The response variable was the number of larvae, while the explanatory variables included blood source, seasonality, gonotrophic cycle and replicate as interaction effects. Additionally, the logarithm of the number of eggs was incorporated as an offset in the model.

Earlier version: Multiple comparisons among treatments were conducted using the Tukey’s honestly significant difference test (HSD Tukey), incorporating the Kramer (1956) correction for unbalanced data, also known as the Tukey-Kramer method. This approach allows to set the family-wise error rate to 0.05. These contrasts were performed using the multcomp package (Hothorn et al., 2008).

Revised version: Multiple comparisons among treatments were conducted using Tukey’s honestly significant difference test (HSD Tukey), incorporating the Kramer (1956) correction for unbalanced data, also known as the Tukey-Kramer method. This approach allows setting the family-wise error rate to 0.05. These comparisons were performed using the emmeans package (Lenth, 2024).

Model validation was assessed using a simulation-based approach to generate randomized quantile residuals to test goodness of fit (Q-Q plot), homoscedasticity, and outliers and influential points (Cook’s distance and Leverage plots). These analyses were conducted with the DHARMa package (Hartig, 2022).

Added in the revised version: Additionally, the statistical power of the models was assessed using the mixedpower package (Kumle et al., 2021), specifically designed to evaluate the power of Generalized Linear Mixed Models (GLMMs). The analysis involved simulating 1,000 iterations with a sample size ranging from 30 to 150 observations per group. The power analysis aimed to achieve a power threshold of 0.8, indicating an 80% chance of detecting true effects if they exist.

Removed in the revised version: Model accuracy via simulations for the fecundity and fertility models

Removed in the revised version: To assess the accuracy of the results, a simulation approach was conducted using a Monte Carlo algorithm to simulate data and model parameters according to the methodology proposed by Carsey & Harden (2013). For this purpose, parameters including coefficients, treatment means and dispersion parameter (referred to as theta, θ), were extracted from the fecundity and fertility models. Initially, a new dataset consisting of 1000 values (sample size of 125 per treatment) was generated. This was achieved using the rnbinom function from the MASS package, specifying the treatment means (computed from model coefficients) and the theta parameter. Subsequently, a negative binomial GLM was fitted to the simulated data, and the expected means for each of the eight treatments were obtained from the model. This simulation process was repeated 1000 times. The ratio between the number of iterations in which the simulated means fell within the observed 95% confidence interval divided by the total number of simulations was computed as an approximation of the robustness of results.

Results

Added in the revised version: A total of 1,162 egg rafts were obtained for analysis, distributed across the eight treatments as follows: chicken:autumn-I: 260, chicken:autumn-II: 174, chicken:summer-I: 162, chicken:summer-II: 88, mouse:autumn-I: 146, mouse:autumn-II: 96, mouse:summer-I: 167, and mouse:summer-II: 69. The full count of egg rafts for each replicate (24 treatments), along with the means for fecundity, fertility, and hatchability, are summarized in Table 1. This sample size fulfilled the statistical power of 0.8 (Figure SM1).

Added in the revised version: In both models—fecundity and fertility—the interaction between blood source and seasonality was statistically significant (fecundity: LRT X² = 5.69, p < 0.05; fertility: LRT X² = 4.37, p < 0.05) (Tables 2,3; Figures 1,2A). Mosquitoes that fed on chicken blood had the highest fecundity (summer: 136.4 eggs/raft, autumn: 141.5 eggs/raft) and fertility (summer: 123.9 larvae/raft, autumn: 126 larvae/raft) in both seasons. In summer, fecundity and fertility were 7% and 11% higher, respectively, in chicken-fed mosquitoes compared to mouse-fed mosquitoes. In autumn, these differences increased to 46% for fecundity and 52% for fertility in favor of chicken-fed mosquitoes. No significant seasonal variation was observed in fecundity and fertility for chicken-fed mosquitoes, but both measures were statistically significant lower by 24% and 25%, respectively, in autumn compared to summer for mouse-fed mosquitoes.

Removed in the revised version: Effect of blood source and seasonality on fecundity and fertility

Removed in the revised version: Two interaction effects were found to be statistically significant in both the fecundity and fertility models (Tables 1-4). The first interaction effect was found between blood source and seasonality (fecundity: LRT X2 = 25.0, p = 7.72 x 10-7; fertility: LRT X2 = 13.82, p = 0.002), meaning that the effect of blood meal source on the number of eggs and larvae per raft varied with season (Tables 2, 3). In general, mosquitoes that fed on avian blood sources exhibited greater fecundity and fertility than those that fed on mammalian blood source, regardless of the season. When treatments were examined by pairwise comparisons, fecundity and fertility were 13% and 18% greater in autumn compared to summer for bird-fed mosquitoes, while they were 17% and 19% lower for mammal-fed mosquitoes (Fig. 1 A, C, Table 1). The second statistically significant interaction was between blood meal source and the gonotrophic cycle (fecundity: LRT X2 = 9.29, p = 0.0023; fertility: LRT X2 = 14.4, p = 0.00073). This indicated that the effect of blood type on the number of eggs and larvae per raft differed between the first and second feeding trial (Tables 2, 3). For mosquitoes that fed on mammalian blood, fecundity and fertility were 29% lower in the second cycle compared to the first. Although the difference was not statistically significant, fecundity and fertility showed a slight increase in the second cycle for bird-fed mosquitoes. During the second cycle, these outcomes were 41% and 40% higher, respectively, for bird-fed mosquitoes compared to mammal-fed mosquitoes (Fig. 1 B, D, Table 1).

Earlier version: Interaction plots showing the marginal effects of blood meal source-seasonality and blood meal source-gonotrophic cycle on (A, B) fecundity and (C, D) fertility. Predicted mean values and its corresponding confidence intervals are shown. Means sharing same letters are not statistically different.

Revised version: Interaction plot showing the effect of blood meal source (chicken or mouse) and seasonality (autumn or summer) on fecundity (eggs/raft). Points represent the predicted marginal mean values and the bars corresponding to their confidence intervals. Means sharing the same letters are not statistically different (p > 0.05).

Added in the revised version: Interaction plots showing the effect of blood meal source (chicken and mouse) and seasonality (autumn and summer) (A) or blood meal source and gonotrophic cycle (first or second) (B) on fertility (eggs/raft) of Culex quinquefasciatus. Points represent the predicted marginal mean values and the bars corresponding to their confidence intervals. Means sharing the same letters are not statistically different (p > 0.05).

Added in the revised version: In the fertility model, there was also a significant interaction between blood source and gonotrophic cycle (LRT X² = 3.98, p < 0.05), indicating that the impact of blood type on larvae production differed between the first and second cycles (Table 3, Figure 2B). Chicken-fed mosquitoes showed no difference in fertility between cycles, while those fed on mouse blood had 25% lower fertility in the second cycle compared to the first. During the second cycle, fertility was 50% higher in chicken-fed mosquitoes compared to mouse-fed mosquitoes.

Added in the revised version: Figure 3 presents the predicted mean values of fecundity and fertility for the 24 treatments, with no differences observed across the three replicates. The variance of random intercepts among treatments was approximately 0.006 for fecundity and 0.007 for fertility, accounting for 6% and 4% of the total variance in the respective models.

Added in the revised version: Dotplots showing the variation of predicted fecundity (eggs/raft) and fertility (larvae/raft) means along the 24 treatments, combination of blood source (chicken or mouse), seasonality (autumn or summer), gonotrophic cycle (first or second) and replicate (1 to 3).

Added in the revised version: There was no statistical difference among treatments and replicates in terms of hatchability, given that no statistical effects of blood source, seasonality or gonotrophic cycle was observed (Table 4). The mean hatchability for all treatments was 0.89, with a range between 0.84 for the mouse:autumn-II (replicate 1) colony and 0.96 for mouse:autumn-I (replicate 1) (Table 1, Figure 4).

Removed in the revised version: Effect of blood source and seasonality on hatchability

Removed in the revised version: The overall hatchability reached 86%, with a range spanning from 76% for MA-II to a peak of 96% for MA-I, as outlined in Table 1 and depicted in Figure 2. The model found a statistically significant effect of blood meal source (LRT X2 = 73.06, p = 0.0012) and an interaction effect between seasonality and gonotrophic cycle (LRT X2 = 41.02, p = 0.0015) on hatching rates (Fig. 3, Table 4). Pairwise comparisons revealed a statistically significant 15% decrease in hatchability for mosquitoes from autumn from the first to the second gonotrophic cycle. Additionally, mosquitoes fed on mammals had greater hatchability than those fed on birds. In summer, hatchability during the first cycle was also higher than during the second cycle. These two trends were not found to be statistically significant.

Earlier version: Boxplots showing the effects of both blood meal source and seasonality on hatchability for the first and second gonotrophic cycle.

Revised version: Boxplots showing the mean hatchability for the 24 treatments, combining replicate (1-3), blood meal source (chicken and mouse), seasonality (summer and autumn) and gonotrophic cycle (first and second), of Culex quinquefasciatus.

Removed in the revised version: Interaction plot showing the marginal effects of seasonality-gonotrophic cycle on hatchability.

Removed in the revised version: Model accuracy via simulations for the fecundity and fertility models

Removed in the revised version: When comparing the simulated means for each of the eight treatments (combinations of blood meal source, seasonality, and gonotrophic cycle) with the means calculated from the observed data a high level of accuracy exceeding 90% was observed. Most of the simulated means fitted within the original 95% confidence interval, with an accuracy rate greater than 98%. The greatest mismatching among simulated and original data were found in the treatment MS-I, with an accuracy of 93% and 95% for the fecundity and fertility model, respectively. Accuracy remained high despite variations in sample sizes and data dispersion (Figs. 4, 5).

Removed in the revised version: Simulated means of fecundity (eggs per raft) of each treatment along 1,000 iterations. Abbreviations. B: avian blood; M: mammalian blood; A: autumn; S: summer; I: first gonotrophic cycle; II: second gonotrophic cycle; Ac: model accuracy.

Removed in the revised version: Simulated means of fertility (eggs per raft) of each treatment along 1,000 iterations. Abbreviations same as Figure 4.

Discussion

Added in the revised version: Our results show there is an interaction effect between blood meal source and seasonality on reproductive outputs of Cx. quinquefasciatus. While our model hypothesized that fecundity and fertility would be higher in bird-fed mosquitoes than in mammal-fed mosquitoes during summer, and the reverse would occur in autumn, our findings revealed the opposite pattern. Regardless the season, fecundity and fertility in Culex quinquefasciatus were higher in bird-fed mosquitoes compared to those fed on mammals. This pattern aligns with those findings from previous studies (Akoh et al., 1993; Richards et al., 2012; Telang & Skinner, 2019). This trend is not unique to Cx. quinquefasciatus but has also been seen in other Culex species (Shroyer & Silvery, 1972) and Aedes mosquitoes (Harrison et al., 2021). The commonly accepted explanation for this pattern suggests that the nutritional quality of avian blood, with its nucleated and larger erythrocytes, provides greater protein and energy, thereby enhancing these reproductive outputs (Alto et al., 2014). However, our findings offer a mixed perspective by highlighting the interaction between blood meal source and seasonality that further might influence fecundity and fertility.

Added in the revised version: The significant interaction observed between blood meal source and seasonality revealed that mammal-fed mosquitoes had lower fecundity and fertility in autumn compared to summer, while bird-fed mosquitoes maintained similar reproductive outputs across both seasons. This finding diverges from the expectation of consistently higher fecundity and fertility in bird-fed mosquitoes and consistently lower fecundity and fertility in mammal-fed mosquitoes. This suggests that environmental factors such as temperature and photoperiod, which have been shown to affect mosquito fitness (Mordecai et al., 2019; Abouzied, 2017; Mogi, 1992), may interact with blood meal source in complex ways. Although the effects of temperature and photoperiod are generally coupled and follow a unimodal trend with an optimal peak at 28°C (Mordecai et al., 2019), our results indicate that the interaction between seasonality and blood source might play a critical role in determining reproductive success.

Removed in the revised version: Our results revealed that there exists an interaction effect of blood meal source and seasonality on reproductive outputs of Cx. quinquefasciatus. Although we expected a higher fecundity and fertility in summer for bird-fed mosquitoes and higher outputs in autumn for mammal-fed mosquitoes, we found the opposite pattern.

Removed in the revised version: Most studies have reported a greater fecundity and fertility in Cx. quinquefasciatus fed on different bird species (Akoh et al.,1993; Richards et al., 2012; Telang & Skinner,2019). This trend is also present in other species of Culex (Shroyer & Silvery, 1972) and Aedes mosquitoes (Harrison et al., 2021). The widely accepted hypothesis that has emerged to explain why avian blood produces a greater fecundity and fertility even in mosquitoes with a known mammal-preferred host choice focuses on the nutritional quality of host blood. Usually, it is addresses that because avian blood possesses nucleated and larger erythrocytes this offers greater quantities of proteins and energy than mammalian blood (Alto et al., 2014). However, this assumption lacks solid evidence due to studies regarding nutritional quality of avian blood and its impact in mosquito fitness are scarce. Actually, some studies have shown that mammalian blood are nutritionally richer than avian blood in terms of total protein content (Pearce 1977, Sant’Anna et al., 2010). Pre-diuresis, a physiological process where red blood cells and proteins are concentrated in the midgut by excreting large amounts of water, has been postulated as a mechanism for explaining the higher fecundity and fertility observed in bird-fed mosquitoes (Erram et al., 2022). Furthermore, avian blood has higher concentrations of isoleucine than mammalian blood, a crucial limiting amino acid necessary for egg production (Greenberg, 1951, Harrison et al., 2021).

Removed in the revised version: The effect of temperature on fitness follows a unimodal relationship with a narrow optimal range. For mosquitoes of genus Culex and Aedes, this optimal band ranges from 20°C to 30°C, with an optimal at 28°C for species such as Cx. quinquefasciatus (Mordecai et al., 2019). Several studies have corroborated this trend, with higher fecundity and fertility towards the optimal of 28°C (Abouzied, 2017; Mogi, 1992), regardless of the blood source (bird or mammal). Regarding photoperiod, the effects are not as evident as temperature since data are somewhat contradictory. While Mogi (1992) found that mid and long day lengths induced greater fecundity, Costanzo et al. (2015) did not find differences of day length on fecundity. Because temperature and photoperiod are always coupled, in general their effects follow the same unimodal trend. Although our experiment was conducted within this optimal range, the highest fecundity and fertility in optimal conditions (summer) was present only in mammal-fed mosquitoes, while higher fitness in bird-fed mosquitoes was detected in autumn. This difference may indicate an interaction of seasonality with blood source.

Removed in the revised version: While our findings may have certain limitations due to moderate sample sizes in some treatments and their reliance solely on laboratory experiments, they verify an interaction effect between blood meal source and seasonality, which had not been previously investigated or documented in any other study. Nevertheless, the precise reasons behind the observed pattern remain inconclusive. The highest levels of fecundity and fertility observed in bird-fed mosquitoes during autumn, when conditions are more challenging, may confer a reproductive advantage, leading to increased offspring production as an adaptive response to these adverse conditions. In contrast, during the more favorable conditions of summer, this pattern may not be as readily discernible. This proposed hypothesis needs further studies to be confirmed.

Earlier version: A second interaction was observed between blood meal source and gonotrophic cycle. We found that fecundity and fertility during the first cycle were similar between mosquitoes fed on chickens and those fed on mammals, and this was also similar with bird-fed mosquitoes from the second gonotrophic cycle, regardless of the season. The only difference we observed was between the fitness of mammal-fed mosquitoes from the second gonotrophic cycle and the earlier mentioned treatments. This pattern is not different from those found in other studies, as the first and second cycles are generally the most productive (Awahmukalah & Brooks, 1985; Christiansen-Jucht et al., 2015). Although we observed a slight increase of fecundity and fertility in bird-fed mosquitoes between cycles, it was statistically non-significant. To our knowledge, only two studies have reported this increasing pattern: an older one by Bennett (1970) and unpublished data from Obholz (personal communication), both conducted with Ae. aegypti. On the contrary, a decrease in fitness between cycles has been reported in other studies for Cx. quinquefasciatus (Richards et al., 2012; Telang & Skinner, 2019) fed upon birds or mammals. This is consistent with our findings regarding mouse-fed mosquitoes; however, because the fecundity and fertility of mouse-fed mosquitoes during the second cycle yielded the fewest egg rafts, this decreasing pattern might also be an effect due to the insufficient sample size.

Revised version: Additionally, our study found a significant interaction between blood meal source and gonotrophic cycle on fertility. For mammal-fed mosquitoes, fertility was significantly lower in the second gonotrophic cycle compared to the first, a pattern commonly reported in other studies where the first and second cycles are generally the most productive (Awahmukalah & Brooks, 1985; Christiansen-Jucht et al., 2015). This declining fertility pattern has been documented in Cx. quinquefasciatus fed on both birds and mammals (Richards et al., 2012; Telang & Skinner, 2019). However, our results showed that this decline in fertility was not present in bird-fed mosquitoes, indicating a potential role of blood meal source in moderating reproductive output across gonotrophic cycles.

Added in the revised version: We also investigated whether hatchability could account for the observed variations in fertility among different treatments. Our results indicate no difference in hatchability among treatments.

Added in the revised version: While our findings corroborate an interaction effect between blood meal source and seasonality, they are not without caveats, as same as other studies. First, our research was conducted under controlled conditions in a laboratory incubator, which implies that our results may differ from those of field studies. Some of the main limitations of this approach include the use of artificial, constant light and temperature rather than a fluctuating environment, and the use of constant humidity, which could affect mosquito performance. However, this is the first study of its kind, and future investigations should aim for more realistic setups. A second limitation, commonly encountered in ecological and behavioral studies, is determining an appropriate sample size for each treatment, as this can impact the ability to detect desired effects or interactions (Taborsky, 2010). To address this issue, we conducted an a priori power analysis to determine the minimum required sample size and repeated the assay multiple times (three replicates). Furthermore, the presence of interaction effects among our explanatory variables suggests that our chosen sample size may have been adequate. Lastly, another aspect of this research that warrants further discussion is the underlying hypothesis that Culex mosquitoes change their feeding habits from birds to mammals. Given that most species are ornithophagic, why is this shift not directed from one bird host to another? In fact, both patterns—bird-to-mammal (Burkett-Cadena et al., 2011) and bird-to-bird (Kent et al., 2009) shifts—are mentioned in the literature. Future studies are aimed to clarify the complex feeding behaviour of Culex mosquitoes regarding seasonal switching, which may depend on several factors, such as habitat, as evidenced in U.S. Culex populations (Caranci, 2010).

Removed in the revised version: In addition to assessing fecundity, we also investigated whether hatchability could account for the variations in fertility observed among different treatments. The only statistically significant trend we identified was a decrease in the percentage of hatched eggs from the first to the second gonotrophic cycle, but this was observed only in autumn mosquitoes. This decline in hatching rate may be attributed to a constant exposure to cooling conditions, due to the challenging environmental conditions during this season. This temperature-related effect aligns with findings from other studies (van der Linde et al., 1990). However, although our data suggest that changes in fertility may be linked to variations in hatching rate, it becomes apparent that fecundity has a more pronounced impact on larval production due to the effects of blood meal source and seasonality than hatching rate. Hence, the influence of hatchability on the number of larvae per raft is rather limited.

Earlier version: In summary, our study investigates the critical issue of host shifting patterns in Culex mosquitoes, which holds significant importance due to their role as bridge vectors between bird and human hosts. Our model suggests that changes in photoperiod and environmental temperature induce a different physiological state that requires the mosquito to feed on a blood meal source that results in higher fitness (increased fecundity and fertility). We predicted that mosquitoes fed on mammals in autumn will perform higher fecundity and fertility than those fed on birds in the same season. Unexpectedly, we observed the opposite pattern. Despite this, our study confirmed a significant interaction between blood meal source and seasonality, suggesting a complex interplay of factors affecting mosquito reproduction and its relationship with host use. Due to the observed opposite trend, we can consider three plausible explaining scenarios. Firstly, it is possible that additional factors, such as genetic mechanisms or variations in midgut microbiota composition, play a role in influencing host use patterns. Secondly, host switching might be explained by factors related to vector biology other than fitness. Lastly, it remains a possibility that Argentinian Cx. quinquefasciatus populations do not exhibit the same host-switching behaviour observed in US populations, and the observed interaction effects may be linked to different mosquito behaviours. Further investigations are warranted to corroborate our findings and gain a deeper understanding of the interaction between mosquito feeding pattern and seasonality. Ultimately, these investigations will provide valuable insights into mosquito-borne disease transmission and forecasting its emergence.

Revised version: In summary, our study investigates the critical feature of host shifting patterns in Culex mosquitoes, which holds significant importance due to their role as bridge vectors between bird and human hosts. Although our study detects a significant interaction between blood meal source and seasonality, we observed the opposite expected trend, suggesting a complex interplay of factors affecting mosquito reproduction and its relationship with host use. We could consider three possible explanations for our results. Firstly, it is possible that additional factors, such as genetic mechanisms or variations in midgut microbiota composition play a role in influencing host use patterns. Secondly, host switching might be explained by factors related to vector biology other than fitness. Lastly, it remains a possibility that Argentinian Cx. quinquefasciatus populations do not exhibit the same host-switching behaviour observed in US populations, and the observed interaction effects may be linked to different mosquito behaviours. However, Spinsanti et al. (2008) found a seasonal variation in WNV human cases, and Beranek (2019) also observed that in autumn, the proportion of blood meals was higher than in summer for Culex quinquefasciatus. Therefore, more in-field studies are needed to assess seasonal host variation in Culex mosquitoes from Argentina.

Added in the revised version: Further investigations are warranted to corroborate our findings and gain a deeper understanding of the interaction between mosquito feeding pattern and seasonality. Last, these investigations will provide valuable insights into mosquito-borne disease transmission and forecasting its emergence.

Added in the revised version: Supporting information